```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

<a id="top"></a>

# Homework 7: IMDb

We will use SQL to dive deep into the Internet Movie Database (IMDb) and answer different questions involving movies, actors, and movie ratings.

## Due Date: Thursday, April 17th, 11:59 PM

You must submit this assignment to Gradescope by the on-time deadline, Thursday, April 17th, 11:59 PM. Please read the syllabus for the Slip Day policy. No late submissions beyond what is outlined in the Slip Day policy will be accepted. **We strongly encourage you to submit your work to Gradescope several hours before the deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, we will not respond to last-minute requests for assistance (TAs need to sleep, after all!).

Please read the instructions carefully when submitting your work to Gradescope. 


##  Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the homework, we ask that you **write your solutions individually**. If you do discuss the assignments with others, please **include their names** below.

**Collaborators**: *list collaborators here*

## Grading 

Grading is broken down into autograded answers and free response. For autograded answers, the results of your code are compared to provided and/or hidden tests. For free response, readers will evaluate how well you answered the question and/or fulfilled the requirements of the question.

<!--
<details>
    <summary>[Click to Expand] <b>Scoring Breakdown</b></summary>-->
|Question| Manual | Points|
|---|---|---|
|1a | No | 2 |
|1b | No | 2 |
|2 | No | 3 |
|3 | No | 3 |
|4a | No | 4 |
|4b | Yes | 2 |
|5 | No | 4 |
|Total | 1 |20 |
</details>

## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: Please only have ONE active SQL database running at once on DataHub, meaning do not have multiple homework/lab/lecture notebooks running at the same time**

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.


### Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Data 100 Debugging Guide](https://ds100.org/debugging-guide/). In this guide, you can find general questions about Jupyter notebooks / Datahub, Gradescope, common SQL errors, and more.

```python
# Run this cell to set up your notebook; no further action is needed
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import plotly.express as px
import seaborn as sns
import sqlalchemy
from pathlib import Path
import sql

plt.style.use('fivethirtyeight') # Use plt.style.available to see more styles
sns.set()
sns.set_context("talk")
np.set_printoptions(threshold=5) # Avoid printing out big matrices
%matplotlib inline
%load_ext sql
```

<br/><br/>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## The IMDB (mini) Dataset

We will explore a miniature version of the [IMDb Dataset](https://www.imdb.com/interfaces/). This is the same dataset that we used for this week's lab.

**Caution: Be careful with large SQL queries!!** You may need to reboot your Jupyter Hub instance if it stops responding. To avoid printing out 100k-sized tables, we've adjusted the display limit to ensure that the tables displayed are truncated to 20 rows (though they may contain more rows in reality).

```python
%config SqlMagic.displaylimit = 20
%config SqlMagic.style = 'DEFAULT'
```

Let's set up the SQL database.

## Question 1 - Setup

Please just run the cells below to set up your SQL database and the autograder.

```python
import duckdb
```

```python
# Run this cell to connect to the database
conn = duckdb.connect()
conn.query("INSTALL sqlite")
```

```python
%sql duckdb:///data/imdbmini.db
```

Let's take a look at the table schemas:

```python
%%sql
-- just run this cell --
SELECT * FROM sqlite_master WHERE type='table';
```

From running the above cell, we see the database has 4 tables: `Name`, `Role`, `Rating`, and `Title`.

<details open>
    <summary>[<b>Click to Expand</b>] See descriptions of each table's schema below. You can also find it in the `schemas.txt` file, which is in the same directory as this notebook. We have only included descriptions for columns that could be of potential use in this homework. </summary>
    
**`Name`** – Contains the following information for names of people.
    
- nconst (BIGINT) - alphanumeric unique identifier of the name/person
- primaryName (VARCHAR) - name by which the person is most often credited
- birthYear (VARCHAR) - in YYYY format
- deathYear (VARCHAR) - in YYYY format
    
    
**`Role`** – Contains the principal cast/crew for titles.
    
- tconst (BIGINT) - alphanumeric unique identifier of the title
- ordering (VARCHAR) - a number to uniquely identify rows for a given tconst
- nconst (BIGINT) - alphanumeric unique identifier of the name/person
- category (VARCHAR) - the category of job that person was in
- characters (VARCHAR) - the name of the character played if applicable, else '\\N'
    
**`Rating`** – Contains the IMDb rating and vote information for titles.
    
- tconst (BIGINT) - alphanumeric unique identifier of the title
- averageRating (VARCHAR) – weighted average of all the individual user ratings
- numVotes (VARCHAR) - number of votes (i.e., ratings) the title has received
    
**`Title`** - Contains the following information for titles.
    
- tconst (BIGINT) - alphanumeric unique identifier of the title
- titleType (VARCHAR) -  the type/format of the title
- primaryTitle (VARCHAR) -  the more popular title / the title used by the producers on promotional materials at the point of release
- isAdult (VARCHAR) - 0: non-adult title; 1: adult title
- startYear (VARCHAR) - represents the release year of a title.
- runtimeMinutes (VARCHAR) - primary runtime of the title, in minutes
    
</details>

<br/><br/>
From the above descriptions, we can conclude the following:
* `Name.nconst` and `Title.tconst` are primary keys of the `Name` and `Title` tables, respectively.
* `Role.nconst` and `Role.tconst` are **foreign keys** that point to `Name.nconst` and `Title.tconst`, respectively.

For a refresher on how to write SQL queries in Jupyter, feel free to reference the start of Lab 10. The main thing to keep in mind is that you can directly write your query in place of the ellipsis under `%%sql --save query_q`. **Please do not edit this line.** 

For example, we can set `query_example` to the output of the following SQL query and directly see what it contains: 

```python
%%sql --save query_example
-- This is a one-line SQL comment.
/* This is a multi-line
   SQL comment. */
SELECT * 
FROM name
LIMIT 15;
```

<br/>

---

### Question 1a  
Let's determine whether our database includes information going back to the early days of cinema or just more recent data.

List the **10 oldest movie titles** by `startYear` and then `primaryTitle` both in **ascending** order. The output should contain the `startYear`, `primaryTitle`, and `titleType`. In this homework, we define a movie as having `titleType='movie'`. Keep this in mind for later questions as well.
                                                                                                                               

```python
%%sql --save query_q1a

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q1a
res_q1a = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q1a")
```

<br/>

---

### Question 1b

Next, let's calculate the distribution of movies by year. Write a query that returns the **total** number of movie titles for each `startYear` in the `Title` table as `total`. Order your final results by the `startYear` in **ascending** order. As in `q1a`, remember that movies are defined as having `titleType=’movie’`.

The first few records of the table should look like the following (but you should compute the entire table):


|startYear|total|
|------:|-----:|
| 1915|1|
| 1920|1|
| 1921|1|
| 1922|1|
| ...|...|



```python
%%sql --save query_q1b

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q1b
res_q1b = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q1b")
```

<br/>

The following cell should generate an interesting plot of the number of movies that premiered each year. Notice there are fewer movies premiering from the 1920s to the late 1940s. Why might that be? *This question is rhetorical; you do not need to write your answer anywhere.*

```python
# Run this call to generate the bar plot; no further action is needed
px.bar(res_q1b, x="startYear", y="total", 
        title="Number of movies premiered each year", 
        width=900, height=400)
```

<br/>

---

## Question 2

Write a SQL query to determine the movie producers with the highest average ratings across all of their movies. The term **"movie actor"** is defined as anyone with an `"actor"` or `"actress"` job category role in a `"movie"` title type released after 1980.

Your SQL query should output exactly two fields named `name` (the movie actor's name) and `total` (the number of movies the movie actor appears in). Order the records by `total` in **descending** order, and break ties by ordering by `name` in **ascending** order.

Your result should look something like this (but without `????`):

| name | total |
|-----:|-----:|
| ???? | 58 |
| ???? | 54 |
| ???? | 53 |
| ???? | 49 |
| ???? | 46 |
| ???? | 43 |
| ???? | 41 |
| ???? | 40 |
| ???? | 40 |
| ???? | 39 |

**Notes**: 
* **The query should take < 2 minutes to run.**
* Sometimes Python gets confused and colors some SQL queries red; *don't worry if the SQL coloring doesn't match what you'd expect*. As long as it runs, it's fine.

**Hints**:

* Before writing your query, you may wish to review the table descriptions given at the start of the assignment to determine where the information you need is stored
* If you want to include a non-aggregate field in the `SELECT` clause, it must also be included in the `GROUP BY` clause.
* When using multiple conditions in a `WHERE` clause, pay attention to the order of operations.

```python
%%sql --save query_q2

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q2
res_q2 = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q2")
```

<br/>

---

## Question 3: The `CASE` Keyword

The `rating` table has the `numVotes` and the `averageRating` for each title. A movie is considered a **"big hit**" if there are more than 100,000 votes for the movie. Which `movie` titles were **"big hits"**? Construct a query that generates the following result:

| isBigHit | total |
|-----:|-----|
| no | ???? |
| yes | ???? |

Where `????` is replaced with the correct values. The row with `no` should have the count for how many movies **are not** big hits, and the row with `yes` should have the count of how many movies **are** big hits.

**Hints**:

* Check the data types of `numVotes` before performing any arithmetic operations.
* You will need to use some type of `JOIN`.
* You may also consider using a `CASE` statement:
    ```
    CASE 
        WHEN ... THEN ...
        ELSE ... 
    END
    ```
    </br>
    
  `CASE` statements are the SQL equivalent of `Python` `if... elif... else` statements. To read up on `CASE`, take a look at the following links:
    - https://mode.com/sql-tutorial/sql-case/
    - https://www.w3schools.com/sql/sql_ref_case.asp


```python
%%sql --save query_q3

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q3
res_q3 = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q3")
```

<br/>

---

## Question 4

Let's compute the average runtime and average rating for `movie` titles by release year and examine how they evolve to see how the movie runtimes and ratings have changed over time.

### Question 4a

We will group movies by their release year (`startYear`). For each year, compute:
- the **average runtime** of all movies (as `avgRuntime`)
- the **average IMDb rating** (as `avgRating`)

Construct a query that generates a table containing the `startYear`, `avgRuntime`, and `avgRating`. Only include movies with **at least 25,000 votes**. Order the final results by the value of `startYear` in ascending order.

**Hints**:  
* Like in previous questions, we are only focusing on **movies**.
* Check the data types of `runtimeMinutes`, `averageRating`, and `startYear` before performing any arithmetic operations — you should `CAST` them to `int`, `float`, and `int` respectively.
* You may define computed columns (e.g., `avgRuntime`) in your `SELECT` clause and reference them in the output.

**Note**: DO NOT cast `averageRating` as an integer. Doing so reduces the precision of the resulting values. See the hint above for what you should cast instead!

```python
%%sql --save query_q4a

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q4a
res_q4a = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

<br>

---

### Question 4b

Create two line plots below. The first should show how the **average movie runtime** has changed over time; the second should show how the **average movie rating** has changed over time. The x-axis should be `startYear` for both plots. Use the columns from the table generated in the previous part, `res_q4a`.

**Notes**:
* Please use `sns` or `plt` functions for plotting.
* Please include descriptive titles and labels.
* If your plot does not show up in the generated PDF, please upload a PDF with a screenshot of your code and the plot.

```python
plt.figure(figsize=(12, 5))

plt.subplot(1, 2, 1)  # DO NOT MODIFY THIS LINE
...

plt.subplot(1, 2, 2)  # DO NOT MODIFY THIS LINE
...

plt.tight_layout()
```

Write 1–2 sentences describing any trends you observe in each plot. This will be graded for completion.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/>

---

## Question 5

Write a SQL query to determine the movie producers with the highest average ratings across all of their movies. Define a **"movie producer"** as anyone with a `producer` job category role in a `movie` title type. Construct a query that generates a table consisting of the **producer's name** (as `name`) and their **average producer rating** (as `producerRating`), computed by rescaling the ratings for movies they produced by the number of votes received by each movie. After rescaling the score for each movie, divide the summation of all rescaled ratings by the total number of votes received by each producer. The formula is below:

$$
\text{producerRating} = 
\frac{\sum_m (\text{averageRating}[m] * \text{numVotes}[m] )}{\sum_m \text{numVotes}[m]}
$$

To make clear the summation, "m" refers to a particular movie that a producer worked on. Thus, the summation over "m" refers to the summation across all movies that a particular producer contributed to.

In addition to the above, only consider ratings where there are **at least 22,500** votes and only consider movie producers that have **at least 20 rated movies**. Present the producers with the **top 10** `producerRating` in **descending** order and break ties alphabetically using the producer's name.

The results should look something like this but without the `????`, and with higher rating precision.

| name            | producerRating |
|-----------------|----------------|
| ???             | 7.76...        |
| ???             | 7.62...        |
| ???             | 7.59...        |
| ???             | 7.43...        |
| ???             | 7.41...        |
| ???             | 7.35...        |
| ???             | 7.30...        |
| ???             | 7.27...        |
| ???             | 7.25...        |
| ???             | 7.24...        |

**Hint**: Check the data types of `numVotes` and `averageRating` before performing any arithmetic operations.

**Notes**:
* ***The query should take < 3 minutes to run.***
* DO NOT cast `averageRating` **as an integer**. Doing so reduces the precision of the resulting values, so your table may not match up exactly with what is shown below.
* If a producer has multiple `role` listings for a movie, then that movie will have a bigger impact on the overall average (this is desired).


```python
%%sql --save query_q5

...
```

```python
# Run this cell for grading purposes. 
# No further action is required. 
query = %sqlcmd snippets query_q5
res_q5 = pd.read_sql(query, "duckdb:///data/imdbmini.db")
```

```python
grader.check("q5")
```


<br><br>


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Chanel and Opal congratulate you for finishing Homework 7!

<img src="images/chanel.jpg" width="240"/> <img src="images/opal.jpg" width="320"/>

### Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly, weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSdpKA_E3b7PGqKSRqBUgSebb9bVFhRwRBv1ueisGsnKFkZkYg/viewform?usp=sf_link). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Once you submit this file to the Homework 7 Coding assignment on Gradescope, Gradescope will automatically submit a PDF file with your written answers to the Homework 7 Written assignment. If you run into any issues when running this cell, feel free to check the ["Why did a Gradescope test fail when all the Jupyter notebook’s tests passed?"](https://ds100.org/debugging-guide/autograder_gradescope/autograder_gradescope.html#why-did-a-gradescope-test-fail-when-all-the-jupyter-notebooks-tests-passed) section of the Data 100 Debugging Guide.

If there are issues with automatically generating the PDF, try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work, download the notebook as an HTML file by clicking on `File -> Save and Export Notebook As... -> HTML`, open the HTML file in your browser, and Ctrl + P or Cmd + P (Mac) to save the file as a PDF. Alternatively, you can  manually take screenshots of your answers to the manually graded questions and submit a PDF containing those. 

**You are responsible for ensuring your submission follows our requirements and that everything was generated and submitted correctly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 2)
```
>>> assert bool(res_q1a.shape == (10, 3) or res_q1a.shape == (15, 3))
```
```
>>> assert bool(set(res_q1a.columns) == set(['startYear', 'primaryTitle', 'titleType']))
```
```
>>> assert bool(np.all(res_q1a['titleType'] == 'movie'))
```
```
>>> actual_years = list(res_q1a['startYear'].astype(str))
>>> expected_sets = [['1902', '1915', '1920', '1921', '1922', '1924', '1925', '1925', '1926', '1927'], ['1915', '1920', '1921', '1922', '1924', '1925', '1925', '1926', '1927', '1927'], ['1915', '1920', '1921', '1922', '1924', '1925', '1925', '1926', '1927', '1927', '1928', '1928', '1929', '1930', '1931']]
>>> assert actual_years in expected_sets
```
```
>>> assert bool(res_q1a.loc[5, 'primaryTitle'] == 'Battleship Potemkin')
```


Tests for q1b (points: 2)
```
>>> assert bool(res_q1b.shape == (102, 2))
```
```
>>> assert bool(res_q1b.size == 204)
```
```
>>> assert bool(np.all(np.isclose(list(res_q1b.median()), [1970.5, 19.0])))
```


Tests for q2 (points: 3)
```
>>> assert bool(res_q2.shape == (10, 2))
```
```
>>> assert bool(set(res_q2.columns) == set(['name', 'total']))
```
```
>>> assert bool(sum(res_q2['total']) == 463)
```
```
>>> assert bool(sorted(list(res_q2['name'])) == ['Adam Sandler', 'Bruce Willis', 'Johnny Depp', 'Liam Neeson', 'Mark Wahlberg', 'Morgan Freeman', 'Nicolas Cage', 'Robert De Niro', 'Samuel L. Jackson', 'Tom Hanks'])
```


Tests for q3 (points: 3)
```
>>> assert bool(res_q3.shape == (2, 2))
```
```
>>> assert bool(set(res_q3.columns) == set(['isBigHit', 'total']))
```
```
>>> assert bool(set(res_q3['isBigHit']) == set(['yes', 'no']))
```
```
>>> assert bool(set(res_q3['total']) == {2041, 4318})
```


Tests for q4a (points: 4)
```
>>> assert set(res_q4a.columns) == {'startYear', 'avgRuntime', 'avgRating'}
```
```
>>> assert res_q4a['avgRuntime'].between(45.0, 152.0).all()
>>> assert res_q4a['avgRating'].between(6.470707, 8.5).all()
```
```
>>> assert res_q4a.dtypes['startYear'] == 'int64'
>>> assert res_q4a.dtypes['avgRuntime'] == 'float64'
>>> assert res_q4a.dtypes['avgRating'] == 'float64'
```
```
>>> assert list(res_q4a['startYear']) == sorted(res_q4a['startYear'])
```
```
>>> assert len(res_q4a) == 100
```


Tests for q5 (points: 4)
```
>>> assert res_q5.shape == (10, 2)
```
```
>>> assert set(res_q5.columns) == set(['name', 'producerRating'])
```
